# 1. feladatsor - Iterátorok
## Névsor

Adott egy tanulói névsor:

`nevek = ["Szandra Eszti", "Zsombor Kristóf", "Ervin Karcsi", "Gabriella Gertrúd ", "Felicitás Alexander", "Aurél Domonkos", "Vilmos Mariska", "Mónika Melánia ", "Henriett Kató", "Dominika Vida", "Felicitás Salamon", "Regina Irma", "Laci Dzsenifer ", "Sándor Judit", "Lilla Bálint", "Terézia Györgyike", "Albert Ábrahám "]`

Készíts egy iterátort `Nevsor` neven, ami ezt a névsort kapja meg, és betűrendben adja vissza a neveket.

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
nevek = ["Szandra Eszti", "Zsombor Kristóf", "Ervin Karcsi", "Gabriella Gertrúd ", "Felicitás Alexander", "Aurél Domonkos", "Vilmos Mariska", "Mónika Melánia ", "Henriett Kató", "Dominika Vida", "Felicitás Salamon", "Regina Irma", "Laci Dzsenifer ", "Sándor Judit", "Lilla Bálint", "Terézia Györgyike", "Albert Ábrahám "]

## BEGIN SOLUTION

### END SOLUTION
for i in Nevsor(nevek):
    print(i)
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
assert_equal([x for x in Nevsor(nevek)], ["Albert Ábrahám ",
"Aurél Domonkos",
"Dominika Vida",
"Ervin Karcsi",
"Felicitás Alexander",
"Felicitás Salamon",
"Gabriella Gertrúd ",
"Henriett Kató",
"Laci Dzsenifer ",
"Lilla Bálint",
"Mónika Melánia ",
"Regina Irma",
"Szandra Eszti",
"Sándor Judit",
"Terézia Györgyike",
"Vilmos Mariska",
"Zsombor Kristóf"])
```

## Négyzetszámok

Készíts egy iterátort `Squares` néven, ami végigmegy egy `limit` ig a számokon 0-tól és mindegyikre visszaadja a négyzetét.

A `limit` legyen `__init__` paraméter.

<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
        
### END SOLUTION
my_squares = Squares(limit=20)
print([x for x in my_squares])
```

# 2. feladatsor - Generátorok
## Hét napjai generátor

Adottak a hét napjai

`days = ["Monday","Tuesday","Wednesday","Thursday","Friday","Saturday","Sunday"]`

Írj egy generátort `day_of_week` néven, ami kigenerálja ezeket.

<!-- grade_id: cell-5f3ac85c68017818, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION

### END SOLUTION
print([x for x in day_of_week()])
```